---
title: Test
description: Test
---

## First exercise

```yaml
type: NormalExercise
key: a0217ef05b
lang: python
xp: 100
skills: 2
```

First assignement

`@instructions`
Import numba and numpy

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}
# import numpy
```

`@solution`
```{python}

```

`@sct`
```{python}

```
